rule_files:
  - keys-api-outdated-validators.rule.yml

evaluation_interval: 1m

tests:
  # Validators registry is disabled
  - interval: 15m
    input_series:
      - series: validators_registry_enabled{}
        values: 0 0 0 0 0
      - series: lido_keys_api_validators_registry_last_update_block_timestamp{}
        values: 0 0 0 0 0
    alert_rule_test:
      - eval_time: 60m
        alertname: KeysApiOutdatedValidators
        exp_alerts: []
  # Outdated validators list
  - interval: 15m
    input_series:
      - series: validators_registry_enabled{}
        values: 0 0 0 0 1 1
      - series: lido_keys_api_validators_registry_last_update_block_timestamp{}
        values: 0 0 0 0 0 900
    alert_rule_test:
      - eval_time: 60m
        alertname: KeysApiOutdatedValidators
        exp_alerts:
          - exp_labels:
              severity: critical
              service: keys_api
              app_team: tooling
            exp_annotations:
              summary: Validators are outdated
              description: Validators were not updated for more than 1h 0m 0s
      - eval_time: 75m
        alertname: KeysApiOutdatedValidators
        exp_alerts:
          - exp_labels:
              severity: critical
              service: keys_api
              app_team: tooling
            exp_annotations:
              summary: Validators are outdated
              description: Validators were not updated for more than 1h 0m 0s

  # Actual validators list
  - interval: 15m
    input_series:
      - series: validators_registry_enabled{}
        values: 1 1 1 1
      - series: lido_keys_api_validators_registry_last_update_block_timestamp{}
        values: 0 0 1 901
    alert_rule_test:
      - eval_time: 60m
        alertname: KeysApiOutdatedValidators
        exp_alerts: []
      - eval_time: 75m
        alertname: KeysApiOutdatedValidators
        exp_alerts: []
